---
title: 'Chapter 3 - Boosted Trees with XGBoost'
description: ""
---

## 3.2.1 Tree Visualization and Feature Importance

```yaml
type: VideoExercise
key: 7614f67a61
xp: 50
```

`@projector_key`
c40c0abdebfacff05ab809473f552810

---

## 3.2.3 Feature Impact on Scorecarding

```yaml
type: NormalExercise
key: bea9fda493
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
