---
title: 'Chapter 1 - Data Preprocessing and Exploration'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## 2.4 Easy Hyperparameters with GridSearch

```yaml
type: VideoExercise
key: 0ccc98e442
xp: 50
```

`@projector_key`
50d2db29b3db8d93542da0d39fcc8cb0
